<img style="float: left;" src="../images/earth-lab-logo-rgb.png" width="150" height="150" />

# Earth Analytics Education

# Week X Homework Template

This week, you will explore `topic X`.  You will use `Y data` provided by `Z source` to complete the following tasks:
1. A
2. B
3. C

```python
# Core imports needed for grading
import matplotlib.pyplot as plt
import matplotcheck.notebook as nb
import matplotcheck.timeseries as ts
```

<!-- grade_id: func-def, points: 0, grade: true, solution: false -->
```python
# These functions will be removed from here and imported to simplify notebook
def run_test(func, points, *args, correct_message="default correct", error_message="default error", **kwargs):
    results = {"points": 0, "pass": False}
    score = 0
    try:
        fname = func.__name__
        results['description'] = fname
        func(*args, **kwargs)
    except Exception as e:
        results['message'] = error_message
        results['traceback'] = e
        pass
    else:
        results['pass'] = True
        results['message'] = correct_message
        results['points'] = points
    
    return results
   
def output_results(results):
    points = 0
    for r in results:
        points += r['points']
        print("Results for test '{}':".format(r['description']))
        if (r['pass']):
            print(" Pass! {msg} ({p} points)".format(msg=r['message'],p=r['points']))
        else:
            print(" Fail! {msg} ({p} points)".format(msg=r['message'],p=r['points']))
            print(" Traceback: {t}".format(t=r['traceback']))
    return points
```

## Import Python Packages

In the cell below, import the necessary `Python` packages that you need to complete this assignment. 

<!-- grade_id: student-imports-answer, points: 2, grade: true, solution: true -->
```python
# Import necessary Python packages

# BEGIN SOLUTION

import os
import pandas as pd
import earthpy as et

# END SOLUTION
```

## Download Data and Set Working Directory

In the cell below, download the necessary data for this assignment and set the working directory. 

<!-- grade_id: student-dir-answer, points: 0, grade: false, solution: true -->
```python
# BEGIN SOLUTION

# Download data 
url = "https://ndownloader.figshare.com/files/12767933"
et.data.get_data(url=url)

# Set working directory
os.chdir(os.path.join(et.io.HOME, 'earth-analytics'))

# END SOLUTION
```

<!-- grade_id: hidden-dir-tests, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

results = []

## Test working directory
def assert_work_dir():
    assert os.getcwd() == os.path.join(et.io.HOME, 'earth-analytics'), "Incorrect working directory"

work_dir_test = run_test(assert_work_dir, 2, 
                  correct_message="Your working directory is correct.",
                  error_message="Check your working directory.")

results.append(work_dir_test)

output_results(results)
    
### END HIDDEN TESTS
```

## Create Plot of X

In the cell below, add your code to create a plot of `X`.  

Be sure to include:
1. An appropriate plot title.
2. Clear labels for the axes (including units of measurements where applicable).

<!-- grade_id: student-plot-answer, points: 0, grade: false, solution: true -->
```python
# Create plot

# BEGIN SOLUTION

# Import to pandas dataframe
precip_2007_to_2017_months = pd.read_csv("data/earthpy-downloads/boulder-precip-2007-to-2017-months-seasons.csv")

fig, ax = plt.subplots()

# Add the x-axis and the y-axis to the plot
ax.plot(precip_2007_to_2017_months.months, precip_2007_to_2017_months.y2013, color="blue")

# Set plot title
ax.set(title="Monthly Precipitation in 2013 for Boulder, CO")

# Add labels to the axes
ax.set(xlabel="Month", ylabel="Precipitation (mm)")
    
# rotate tick marks on x-axis
plt.setp(ax.get_xticklabels(), rotation=45);

# END SOLUTION

### DO NOT REMOVE LINE BELOW ###
ts_1_plot = nb.convert_axes(plt)
```

## Explain Plot

In the `Markdown` cell below, answer the following questions about your plot:
1. A?
2. B?

<!-- grade_id: student-explain-plot-answer, points: 7, grade: true, solution: true -->
### Replace this text with your answer


## Data Questions

In the cell below, add your code to determine the following about `Y data`:
1. A
2. B

```python
# BEGIN SOLUTION

precip_2007_to_2017_months.describe()

# END SOLUTION
```

<!-- grade_id: visible-plot-tests, points: 6, grade: true, solution: false -->
```python
# Locked Grading Cell
## Visible test
hw_plot1 = ts.TimeSeriesTester(ts_1_plot)

plot_x_lab_test = run_test(hw_plot1.assert_axis_label_contains, 2, axis="x", lst=["Month"],
                  correct_message="Your x axis label contains the appropriate keywords.",
                  error_message="Check your x axis label for a more appropriate label.")
results.append(plot_x_lab_test)

plot_y_lab_test = run_test(hw_plot1.assert_axis_label_contains, 2, axis="y", lst=["Precip", "mm"], 
                  correct_message="Your y axis label contains the appropriate keywords and units identified.",
                  error_message="Check your y axis label for a more appropriate label with units identified.")
results.append(plot_y_lab_test)

plot_title_test = run_test(hw_plot1.assert_title_contains, 2, ["Month", "Precip", "2013"],
                   correct_message="Plot title contains the appropriate keywords.",
                   error_message="Check your plot title words for a more appropriate title.")
results.append(plot_title_test)

output_results(results)
```

<!-- grade_id: hidden-plot-tests, points: 9, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

plot_type_test = run_test(hw_plot1.assert_plot_type, 3, 'line', 
                 correct_message="Your plot is of the correct type.",
                 error_message="Check your plot type.")
results.append(plot_type_test)

# Create expected data: full correct answer
precip_2007_to_2017_months = pd.read_csv("data/earthpy-downloads/boulder-precip-2007-to-2017-months-seasons.csv")
precip_2007_to_2017_months.reset_index(inplace=True)
precip_2013 = precip_2007_to_2017_months[["index", "y2013"]]

# Add test to ensure the df has the correct values
data_test = run_test(hw_plot1.assert_xydata, 6, xy_expected=precip_2013, 
                correct_message="Plot includes correct data",
                error_message="Check your plot data.",
                xcol="index",
                ycol="y2013")

results.append(data_test)

output_results(results)

### END HIDDEN TESTS
```